---
title: 'Econometrics: Wage gap'
description: 'This chapter deals with estimating the gender pay gap. '
---

## Analyzing the Gender Pay Gap

```yaml
type: VideoExercise
key: b1e3055583
xp: 50
```

`@projector_key`
c8de6b3084beec2c4e24041a02e08d88

---

## Ex 1.1: Introducing the data set

```yaml
type: NormalExercise
key: 42ecbeba83
lang: r
xp: 100
skills: 1
```

The goal of this exercise is to

- load the dataset and tidy it
- get an overview of the variables in the dataset
- print summary statistics
- plot the distribution 

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. It should get students 50% of the way to the correct answer.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
